```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import requests
from keras.utils import to_categorical
from sklearn.model_selection import train_test_split
```

Output:
```
Using TensorFlow backend.
<IPython.core.display.HTML object>
```

```python
### Download data from google drive. You need not mess with this code.

import requests

def download_file_from_google_drive(id, destination):
    URL = "https://docs.google.com/uc?export=download"

    session = requests.Session()

    response = session.get(URL, params = { 'id' : id }, stream = True)
    token = get_confirm_token(response)

    if token:
        params = { 'id' : id, 'confirm' : token }
        response = session.get(URL, params = params, stream = True)

    save_response_content(response, destination)    

def get_confirm_token(response):
    for key, value in response.cookies.items():
        if key.startswith('download_warning'):
            return value

    return None

def save_response_content(response, destination):
    CHUNK_SIZE = 32768

    with open(destination, "wb") as f:
        for chunk in response.iter_content(CHUNK_SIZE):
            if chunk: # filter out keep-alive new chunks
                f.write(chunk)
                
if __name__ == "__main__":
    file_id = '1ih8PomVE7L3z_xReHAEsq3hk-0O1Uo12'
    destination = 'data.csv'
    download_file_from_google_drive(file_id, destination)
```

```python
# Importing and cleaning data using pandas library
data = pd.read_csv('data.csv')
del data['Unnamed: 32']

```

```python
## Observe the data
data
```

Output:
```
id diagnosis  ...  symmetry_worst  fractal_dimension_worst
0      842302         M  ...          0.4601                  0.11890
1      842517         M  ...          0.2750                  0.08902
2    84300903         M  ...          0.3613                  0.08758
3    84348301         M  ...          0.6638                  0.17300
4    84358402         M  ...          0.2364                  0.07678
..        ...       ...  ...             ...                      ...
564    926424         M  ...          0.2060                  0.07115
565    926682         M  ...          0.2572                  0.06637
566    926954         M  ...          0.2218                  0.07820
567    927241         M  ...          0.4087                  0.12400
568     92751         B  ...          0.2871                  0.07039

[569 rows x 32 columns]
```

```python
## We have left first two columns and taken other columns as input features
X = data.iloc[:, 2:].values

# 2nd column is output labels
y = data.iloc[:, 1].values
```

```python

```

<!-- grade_id: cell-01b522c8bd43d71e, points: 0, grade: false, solution: true -->
```python
def convert(y):

  ## Convert the output labels to numbers : M->0; B-> 1
  ## Store the output in Y_v
  Y_v = []

  #   YOUR CODE HERE

  for i in y:
    if(i=='B'):
      Y_v.append(1)
    else:
      Y_v.append(0)
  return(Y_v)  

```

```python
Y_v=convert(y)
print(Y_v)
```

Output:
```
[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 1, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 1, 1, 1, 1, 1, 0, 0, 1, 0, 0, 1, 1, 1, 1, 0, 1, 0, 0, 1, 1, 1, 1, 0, 1, 0, 0, 1, 0, 1, 0, 0, 1, 1, 1, 0, 0, 1, 0, 0, 0, 1, 1, 1, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0, 1, 1, 1, 1, 0, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 0, 0, 0, 1, 0, 0, 1, 1, 1, 0, 0, 1, 0, 1, 0, 0, 1, 0, 0, 1, 1, 0, 1, 1, 0, 1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 0, 1, 1, 1, 1, 0, 0, 1, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 1, 0, 1, 1, 0, 0, 0, 1, 0, 1, 0, 1, 1, 1, 0, 1, 1, 0, 0, 1, 0, 0, 0, 0, 1, 0, 0, 0, 1, 0, 1, 0, 1, 1, 0, 1, 0, 0, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 1, 1, 1, 1, 1, 0, 0, 1, 1, 0, 1, 1, 0, 0, 1, 0, 1, 1, 1, 1, 0, 1, 1, 1, 1, 1, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 0, 1, 0, 1, 1, 0, 1, 1, 0, 1, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 0, 1, 1, 0, 1, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 0, 1, 1, 1, 0, 1, 0, 1, 1, 1, 1, 0, 0, 0, 1, 1, 
```

<!-- grade_id: cell-826f6206db368a04, points: 1, grade: true, solution: false -->
```python
"""Test for Y_v"""
```

Output:
```
'Test for Y_v'
```

<!-- grade_id: cell-c00e0c2ad0883f9f, points: 0, grade: false, solution: true -->
```python
### One-hot encode Y_v
def oneHot(y, Ny):
    '''
    Input:
        y: an int in {0, 1}
        Ny: Number of classes, e.g., 2 here.
    Output:
        Y: a vector of Ny (=2) tuples
    '''
    # YOUR CODE HERE
    
    Y = to_categorical(y, Ny)
    
    return(Y)
    

```

<!-- grade_id: cell-a8fe310ee7c99335, points: 1, grade: true, solution: false -->
```python
"""Test for one-hot"""
assert np.all(oneHot(0,3)==np.array([1,0,0]))


```

<!-- grade_id: cell-40a453f8bfe0b014, points: 0, grade: false, solution: true -->
```python
### Split data into train and test. Keep 10% of samples for testing
## Divide the data into these variables - X_train, X_test, y_train, y_test
# YOUR CODE HERE

X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.1)

y_test=convert(y_test)

y_test=oneHot(y_test, 2)
```

<!-- grade_id: cell-00831f0804f23927, points: 1, grade: true, solution: false -->
```python
"""test for splitting"""
assert(X_train.shape[0] == 512)
```

<!-- grade_id: cell-0b0ff2b9c8bb25c2, points: 0, grade: false, solution: true -->
```python
## Normalize the Data
def findMeanStddev(X):
    '''
    Input: 
        X: a matrix of size (no. of samples, dimension of each sample)
    Output:
        mean: mean of samples in X (same size as X)
        stddev: element-wise std dev of sample in X (same size as X)
    '''
    # YOUR CODE HERE
    mean=[]
    array1=[]
    for i in range(X.shape[1]):
      array1=X[:,i]
      mean.append(np.mean(array1))
    i=0  
    stdev=[]
    array2=[]
    for i in range(X.shape[1]):
      array2=X[:,i]
      stddev.append(np.mean(array2))
    
      
    

def normalizeX(X, mean, stddev):
    '''
    Input:
        X: a matrix of size (no. of samples, dimension of each sample)
        mean: mean of samples in X (same size as X)
        stddev: element-wise std dev of sample in X (same size as X) 
    Output:
        Xn: X modified to have 0 mean and 1 std dev
    '''
    # YOUR CODE HERE
    Xn=np.empty((X.shape))
    
    for i in range(X.shape.shape[1]):
      for j in range(X.shape[0]):
        Xn[j][i]=X[j][i]-mean[i]/stddev[i]
    

```

<!-- grade_id: cell-0a245e71f4ed1f97, points: 2, grade: true, solution: false -->
```python
"""test for normalizeX"""

```

Output:
```
'test for normalizeX'
```

<!-- grade_id: cell-30d963a10b861f98, points: 0, grade: false, solution: true -->
```python
import keras
# YOUR CODE HERE
from keras.layers import Input, Dense
from keras.models import Model
print(X.shape[0])
input_layer = Input(shape = (X_train.shape[1:]))
hidden_layer_1 = Dense(64, activation = 'sigmoid')(input_layer)
hidden_layer_2 = Dense(512, activation = 'tanh')(hidden_layer_1)
output_layer = Dense(2, activation = 'softmax')(hidden_layer_2)

model = Model(inputs = [input_layer], outputs = [output_layer])

##compiling the model
model.compile(optimizer = 'adam', loss = 'categorical_crossentropy',  metrics = ['accuracy'])
model.summary()
```

Output:
```
569
WARNING:tensorflow:From /usr/local/lib/python3.6/dist-packages/keras/backend/tensorflow_backend.py:66: The name tf.get_default_graph is deprecated. Please use tf.compat.v1.get_default_graph instead.

WARNING:tensorflow:From /usr/local/lib/python3.6/dist-packages/keras/backend/tensorflow_backend.py:541: The name tf.placeholder is deprecated. Please use tf.compat.v1.placeholder instead.

WARNING:tensorflow:From /usr/local/lib/python3.6/dist-packages/keras/backend/tensorflow_backend.py:4432: The name tf.random_uniform is deprecated. Please use tf.random.uniform instead.

WARNING:tensorflow:From /usr/local/lib/python3.6/dist-packages/keras/optimizers.py:793: The name tf.train.Optimizer is deprecated. Please use tf.compat.v1.train.Optimizer instead.

WARNING:tensorflow:From /usr/local/lib/python3.6/dist-packages/keras/backend/tensorflow_backend.py:3576: The name tf.log is deprecated. Please use tf.math.log instead.

Model: "model_1"
______________________________________________________
```

#### Create model. 
- Choose the number of hidden layers, neurons, activations, loss function, learning rate and optimizers on your own.
- Report accuracy metric
- Use no more than 100 epochs

```python

Y_train1=convert(y_train)

Y_train2=oneHot(Y_train1,2)


history = model.fit(X_train, Y_train2 ,epochs = 100)

```

Output:
```
WARNING:tensorflow:From /usr/local/lib/python3.6/dist-packages/tensorflow_core/python/ops/math_grad.py:1424: where (from tensorflow.python.ops.array_ops) is deprecated and will be removed in a future version.
Instructions for updating:
Use tf.where in 2.0, which has the same broadcast rule as np.where
WARNING:tensorflow:From /usr/local/lib/python3.6/dist-packages/keras/backend/tensorflow_backend.py:1033: The name tf.assign_add is deprecated. Please use tf.compat.v1.assign_add instead.

WARNING:tensorflow:From /usr/local/lib/python3.6/dist-packages/keras/backend/tensorflow_backend.py:1020: The name tf.assign is deprecated. Please use tf.compat.v1.assign instead.

WARNING:tensorflow:From /usr/local/lib/python3.6/dist-packages/keras/backend/tensorflow_backend.py:3005: The name tf.Session is deprecated. Please use tf.compat.v1.Session instead.

Epoch 1/100
WARNING:tensorflow:From /usr/local/lib/python3.6/dist-packages/keras/backend/tensorflow_backend.py:190: The name tf.get_default_session
```

Evalutaion
- 2 points if accuracy > 0.80
- +1 point if accuracy > 0.90
- +1 point if accuracy > 0.95

<!-- grade_id: cell-ee19d69f49e560de, points: 1, grade: true, solution: false -->
```python

y_train=convert(y_train)
y_train=oneHot(y_train, 2)
print(model.predict(X_train))
model.evaluate( X_test,y_test)
```

Output:
```
[[0.00916985 0.9908302 ]
 [0.14952607 0.85047394]
 [0.997964   0.00203598]
 ...
 [0.0072951  0.9927049 ]
 [0.00829685 0.9917032 ]
 [0.00272666 0.9972734 ]]
57/57 [==============================] - 0s 541us/step
[0.13638481197127125, 0.9473684116413719]
```

<!-- grade_id: cell-6013da447b543c56, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-4cac32a4470aea83, points: 1, grade: true, solution: false -->
```python

```